## Week 6 Assignment - W200 MIDS Python

Assignment due date: 11:59PM PST the night before the week 7 live session.

# Part 1 Coding practice **reading and response**

Your task for this week is to write a 250-word reading response to the article below. In addition, please list **3 questions** that you have from the article. Please write your response in a markdown cell, below this cell in the notebook.

The writing response is a free response, so you may write about your reactions, an interesting thing that you saw in the article, something that really stuck out to you, etc.

**Article**: [The PEP 8 Style Guide](https://www.python.org/dev/peps/pep-0008/#overriding-principle). This document is really important for Python coders because it describes best practices and customs for how one should write Python code. Please read it all and prepare your questions for class.

<!-- grade_id: cell-c4a1546ee8708c36, points: 20, grade: true, solution: true -->
#  PEP 8 Style Guide
-Praba Santhanakrishnan

First of all, it is very interesting to see so many people contributed to Python Enhancement Proposals.  Some great observations including the guidelines for importing. Especially the order in which the imports are grouped (Standard, Related Third Party and the local/Library specific imports). After reading this, I got the clarity on what type of imports to do. (Absolute imports are recommended) , I think it might have performance improvements to because of loading only the absolute packages that are required. 

Good point on updating the comment when the code changes, this has been an historic issue when it comes to debugging or investigating a production bug when the comments do not reflect the actual code and mislead the developers.
Interesting observation - Python guidelines recommend using “_” (underscore) a lot more than in any other languages in the naming conventions.  Also, important to note the differences between Python2 versus Python3. (For example: raise ValueError, ‘message’ which is not legal in Python 3

Great to understand the viewpoint on writing code that doesn’t break for different Python implementations. I always used = operator for comparing singletons like None instead of recommended is or is not. Also good suggestion on not relying on the except: for everything in the exception handlers and catching using the types of exceptions as much as possible. Naturally everyone tends to pack as much code in the try clause so that you don’t have to too many try clauses and catch as many exceptions in just a few try clauses. Good recommendation on avoiding these practices.


1. Is there an editor that can warn the user if the PEP 8 Style Guide is not followed?
2. What is the preference or recommendations on when to use dunder names or not?
3. What is the recommendation/Guidelines on creating pacakages and its conventions?





# Part 2  - Cheating at Scrabble

Write a Python script that takes a Scrabble rack **as a command-line argument** and prints all valid Scrabble words that can be constructed from that rack, along with their Scrabble scores, sorted by score. Valid Scrabble words are provided in the data source below.  A Scrabble rack is made up of any 7 characters.


## There are a few of requirements:

- This needs to be able to be **run as a command line tool** as you'll see below.


- We will use the code embedded in the notebook to test your code, if you do the extra credit give us a similar line to test the extra credit functionality

- Put comments in your code and docstrings in the functions as described in the PEP8 Style Guide above.

- Please include a function called `score_word` in a separate module.  Import this function into your main solution code.

- You need to handle input errors from the user and suggest what that error might be caused by (helpful error messages).

- Implement wildcards as either `*` or `?`. That is, let the user specify a wildcard character that can take any value. There can be a total of two wild cards in any user input (one of each character).

Runtime - should run in a reasonable amount of time (single digit minutes or seconds)
You can use a pre-processed dictionary if this makes sense to you


## Extra Credit (+10 points):

- Allow a user to specify that a certain letter has to be at a certain location. Your program must work without it so this should be completely optional. For the extra credit, locations of certain letters must be specified at the command line, **it may not be some sort of user prompt.**


## Example 
An example invocation and output:

```sh
$ python scrabble.py ZAEFIEE
17 feeze
17 feaze
16 faze
15 fiz
15 fez
12 zee
12 zea
11 za
6 fie
6 fee
6 fae
5 if
5 fe
5 fa
5 ef
2 ee
2 ea
2 ai
2 ae
```


## The Data
http://courses.cms.caltech.edu/cs11/material/advjava/lab1/sowpods.zip contains all words in the official SOWPODS word list, one word per line. You should download the word file and keep it in your repository so that the program is standalone (instead of accessing it over the web from Python).

You can read data from a text file with the following code:
```python
with open("sowpods.txt","r") as infile:
    raw_input = infile.readlines()
    data = [datum.strip('\n') for datum in raw_input]
print(data[0:6])
``` 



Here is a dictionary containing all letters and their Scrabble values:
```python
scores = {"a": 1, "c": 3, "b": 3, "e": 1, "d": 2, "g": 2,
         "f": 4, "i": 1, "h": 4, "k": 5, "j": 8, "m": 3,
         "l": 1, "o": 1, "n": 1, "q": 10, "p": 3, "s": 1,
         "r": 1, "u": 1, "t": 1, "w": 4, "v": 4, "y": 4,
         "x": 8, "z": 10}
``` 

## Tips
We recommend that you work on this and try to break down the problems into steps on your own before writing any code. Once you've scoped generally what you want to do, start writing some code and if you get stuck, take a step back and go back to thinking about the problem rather than trying to fix lots of errors at the code level. You should only use the Python standard library in this assignment, however any tool in the Python standard library is fair game.

## Tips Part 2
If you keep getting stuck check out: https://openhatch.org/wiki/Scrabble_challenge. This is where we got the idea for this assignment and it provides some helpful tips for guiding you along the way. If that link doesn't work you can use the [google cached version here](http://webcache.googleusercontent.com/search?q=cache:https://openhatch.org/wiki/Scrabble_challenge). However, we would sincerely recommend that you try to implement this first yourself before looking at the hints on the website.

Good luck!
 

### The code below will test your command line implementation of the scrabble.py code. Some of these tests are available for you to try.


```python

import subprocess
from nose.tools import assert_equal 
from nose.tools import assert_true
from nose.tools import assert_greater
from nose.tools import assert_less

```

<!-- grade_id: cell-6cb374b4f427be46, points: 20, grade: true, solution: true -->
```python
""" This code should tell us we need a rack and there should be reasonable error anticipation """
# no rack error
!python scrabble.py  

### BEGIN HIDDEN TESTS 
### END HIDDEN TESTS
```

Output:
```
Please pass valid argument.
```

<!-- grade_id: cell-679eb0dae15dccd4, points: 0, grade: true, solution: true -->
```python
""" Does not fail due to trivial mistakes and takes wildcards """ 
# does not fail due to case
!python scrabble.py PENguin  
# takes wildcards
!python scrabble.py PEN*?in    

# Word list for 'PENguin' should look like the words_list1 below (formatting might be different)
word_list1 = [('PENGUIN', 10), ('PENING', 9), ('GENIP', 8), ('UNPEG', 8), ('INGENU', 7), ('PENNI', 7), ('PING', 7), 
             ('PUNG', 7), ('UNPEN', 7), ('UNPIN', 7), ('GIP', 6), ('GUP', 6), ('PEG', 6), ('PEIN', 6), ('PENI', 6), 
             ('PIG', 6), ('PINE', 6), ('PUG', 6), ('ENNUI', 5), ('GENU', 5), ('GIEN', 5), ('GINN', 5), ('NEP', 5), 
             ('NIP', 5), ('PEN', 5), ('PIE', 5), ('PIN', 5), ('PIU', 5), ('PUN', 5), ('ENG', 4), ('GEN', 4), ('GIE', 4), 
             ('GIN', 4), ('GNU', 4), ('GUE', 4), ('GUN', 4), ('NEG', 4), ('NINE', 4), ('PE', 4), ('PI', 4), ('UP', 4), 
             ('GI', 3), ('GU', 3), ('INN', 3), ('NIE', 3), ('NUN', 3), ('UG', 3), ('UNI', 3), ('EN', 2), ('IN', 2), 
             ('NE', 2), ('NU', 2), ('UN', 2)]

# Word list for 'PEN*?in is far too long!

### BEGIN HIDDEN TESTS 
### END HIDDEN TESTS
```

Output:
```
[('PENGUIN', 10), ('PENING', 9), ('GENIP', 8), ('UNPEG', 8), ('INGENU', 7), ('PENNI', 7), ('PING', 7), ('PUNG', 7), ('UNPEN', 7), ('UNPIN', 7), ('GIP', 6), ('GUP', 6), ('PEG', 6), ('PEIN', 6), ('PENI', 6), ('PIG', 6), ('PINE', 6), ('PUG', 6), ('ENNUI', 5), ('GENU', 5), ('GIEN', 5), ('GINN', 5), ('NEP', 5), ('NIP', 5), ('PEN', 5), ('PIE', 5), ('PIN', 5), ('PIU', 5), ('PUN', 5), ('ENG', 4), ('GEN', 4), ('GIE', 4), ('GIN', 4), ('GNU', 4), ('GUE', 4), ('GUN', 4), ('NEG', 4), ('NINE', 4), ('PE', 4), ('PI', 4), ('UP', 4), ('GI', 3), ('GU', 3), ('INN', 3), ('NIE', 3), ('NUN', 3), ('UG', 3), ('UNI', 3), ('EN', 2), ('IN', 2), ('NE', 2), ('NU', 2), ('UN', 2)]
[('ZIZ', 21), ('JIZ', 19), ('ZEX', 19), ('PHENIX', 18), ('PHIZ', 18), ('JINX', 18), ('ZINKE', 18), ('BENZIN', 17), ('PIZED', 17), ('HIZEN', 17), ('WINZE', 17), ('WIZEN', 17), ('PYXIE', 17), ('ZINEB', 16), ('PEIZE', 16), ('ZEK', 16), ('MIZEN', 16), ('PIEZO', 16), ('EQUIP', 16), ('PIQUE', 16), ('PRIZE', 16), ('PIZES', 16), ('UNZIP', 16), ('PI
```

<!-- grade_id: cell-c75db1883eda598b, points: 20, grade: true, solution: false -->
```python
""" The code checks if a list of words is created - if this cell runs without an error than it passed """

#test the code in the command line
cmd = [ 'python', 'scrabble.py', 'PENGUIN' ]
out=bytes.decode(subprocess.Popen( cmd, stdout=subprocess.PIPE ).communicate()[0]).strip()
assert_greater(len(out), 0)

### BEGIN HIDDEN TESTS 
### END HIDDEN TESTS
```

<!-- grade_id: cell-bbdac153d30cb686, points: 20, grade: true, solution: false -->
```python
""" The code checks if a list of words is created - if this cell runs without an error than it passed """

cmd = [ 'python', 'scrabble.py', 'PENGU*?' ]
out=bytes.decode(subprocess.Popen( cmd, stdout=subprocess.PIPE ).communicate()[0])
assert_greater(len(out), 0)

### BEGIN HIDDEN TESTS 
### END HIDDEN TESTS

```

<!-- grade_id: cell-d558229c3d138228, points: 20, grade: true, solution: false -->
```python
""" The code should run in seconds to a few minutes. If < 5 minutes this cell will just print out your time"""
import time
start=time.time()

#test the code in the command line
cmd = [ 'python', 'scrabble.py', 'PENGU*?' ]
out=bytes.decode(subprocess.Popen( cmd, stdout=subprocess.PIPE ).communicate()[0])

tot_time=time.time()-start
print('Total time was {} seconds'.format(tot_time))
assert_less(tot_time, 300)

```

Output:
```
Total time was 22.2277193069458 seconds
```

<!-- grade_id: cell-a4d25cc447ae3e6a, points: 10, grade: true, solution: true -->
```python
""" Implement extra credit call/run the code here """
""" The code checks if a list of words is created - if this cell runs without an error than it passed """

!python scrabble.py PENguin P 2

#test the code in the command line
cmd = [ 'python', 'scrabble.py', 'PENGUIN' , 'P','2']
out=bytes.decode(subprocess.Popen( cmd, stdout=subprocess.PIPE ).communicate()[0]).strip()
assert_greater(len(out), 0)

### BEGIN HIDDEN TESTS 
### END HIDDEN TESTS
```

Output:
```
[('UNPEG', 8), ('UNPEN', 7), ('UNPIN', 7), ('GIP', 6), ('GUP', 6), ('NEP', 5), ('NIP', 5)]
```

```python

```